<div class="alert alert-block alert-info">
    
### Human-Centered Interaction and Collaboration in Robotics: Assignment 5

## General Information:
Please do not add or delete any cells. Answers belong into the corresponding cells (below the question). If a function is given (either as a signature or a full function), you should not change the name, arguments or return value of the function.

If you encounter empty cells underneath the answer that can not be edited, please ignore them, they are for testing purposes.

When editing an assignment there can be the case that there are variables in the kernel. To make sure your assignment works, please restart the kernel and run all cells before submitting (e.g. via Kernel -> Restart & Run All).

Code cells where you are supposed to give your answer often include the line `raise NotImplementedError`. This makes it easier to automatically grade answers. If you edit the cell please outcomment or delete this line.

## Submission:
Please submit your notebook via the web interface (in the main view -> Assignments -> Submit). For the deadline please refer to LEA.

## Group Work:
Please enter your UID (your username here) and those of your group partners into the next cell. We apply plagiarism checking, so do not submit others solutions! If an assignment has a copied solution, the task will be graded with 0 points for all people with the same solution.

## Questions about the Assignment:
If you have questions about the assignment please post them in the LEA forum ahead of the deadline. Don't wait until the last day to post questions!
    
</div>

### Please add the usernames of all your team members in the manner member1, member2 in next cell (example given below)

member1 = 'exampleone2s'

member2 = 'exampletwo2s'

member3 = 'examplefour2s'

#### If you are not working in a group or in a smaller group, then please add member3 and/or member2 as none2s

<!-- grade_id: cell-a9c2f4587c671f98, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e27ced94e387fe18, points: 0, grade: true, solution: false -->
```python
# Execute this cell to make sure you correctly filled in the usernames of the team members

def group_name_test():
    for member_id in [member1, member2, member3]:
        assert isinstance(member_id, str), "Please give your member id as a string."
        assert len(member_id) > 0, "You need to fill in the member id for both members"
        assert member_id.endswith("2s"), "The member id should end with 2s (Your JupyterHub username)"

group_name_test() 
print("All tests passed!")
```

# Task-1

### Designing a Human-Robot Interaction Experiment (45%)

Suppose that you are developing a scenario where the Atlas robot should assist a plumber by bringing necessary tools and objects when requested. Let’s say, you have developed several ways in which the robot can handover an object to the plumber. Now, you would like to evaluate these by conducting an empirical study with experienced plumbers. (You can watch this video for inspiration: https://www.youtube.com/watch?v=-e1_QhJ1EhQ)

*Please refer to the following paper to understand the characteristics of good research questions and hypotheses*

**G. Hoffman and X. Zhao, “A Primer for Conducting Experiments in Human–Robot Interaction,” ACM Transactions on Human-Robot Interaction, vol. 10, no. 1, Art. no. 6, pp. 1–31, Oct. 2020, doi: 10.1145/3412374.**

#### Define at least one research question that you would like to investigate in the context of human-robot interaction in the above scenario.

<!-- grade_id: Task-1_A, points: 10, grade: true, solution: true -->
How does the robot’s handover strategy (basic, communicative, adaptive) affect task performance, perceived safety, and trust of experienced plumbers during a tool-delivery task with the Atlas robot?

#### Identify the constructs involved in this research question(s).

<!-- grade_id: Task-1_B, points: 5, grade: true, solution: true -->
Robot Behavior (Independent Construct)
- Handover strategy
- Basic reactive handover
- Communicative handover
- Adaptive handover

Human Experience (Dependent Constructs)
- Task performance
- Perceived safety
- Trust in robot
- Cognitive workload
- User preference / satisfaction

#### Formulate at least one hypothesis based on the research question(s) and the identified constructs.

<!-- grade_id: Task-1_C, points: 10, grade: true, solution: true -->
H1 (Performance)
- Plumbers will complete tool handovers faster with the adaptive handover strategy compared to the basic reactive strategy.

H2 (Safety)
- Perceived safety will be higher in the communicative and adaptive conditions than in the basic condition.

H3 (Trust)
- Trust in the robot will be highest in the adaptive handover condition.

#### How would you operationalise the constructs in the above hypotheses?

<!-- grade_id: Task-1_D, points: 5, grade: true, solution: true -->
1. Task Performance
    - Handover time (seconds)
    - Total task completion time
    - Number of failed handovers or drops
2. Perceived Safety
    - Questionnaire rating (e.g., 1–7 Likert scale)
    - Example item: “I felt safe interacting with the robot”
3. Trust
    - Standard HRI trust scale (Likert-based)
    - Example item: “I trusted the robot to deliver tools correctly”
4. Cognitive Workload
    - NASA-TLX questionnaire
    - Measures mental demand, effort, frustration
5. User Preference
    - Ranking of conditions after experiment
    - Post-study interview

#### How would you design your study (context, conditions, allocation of subjects)? Justify your study design.

<!-- grade_id: Task-1_E, points: 10, grade: true, solution: true -->
Context
- Simulated plumbing workshop environment
- Realistic sink/pipe setup
- Atlas robot acts as assistant delivering toolsConditions (Independent Variable)

Within-subject comparison of 3 handover strategies:

Basic reactive
- Robot simply gives object when asked
Communicative
- Robot speaks + gives cues + announces tool
Adaptive
- Robot adjusts posture, timing, and orientation to user

Participants
- 20–30 experienced plumbers
- Professional background required

Allocation of Subjects
- Design type:
    - Each participant experiences all 3 conditions.

- Why this design?
    - Plumbers are a specialized population so hard to recruit many
    - Reduces variability between individuals
    - Improves statistical power
    - Allows direct comparison of conditions per person

#### What would your null hypothesis look like, for each of the hypotheses that you formulated above?

<!-- grade_id: Task-1_F, points: 5, grade: true, solution: true -->
H0₁ (Performance)
- There is no significant difference in task completion time between basic, communicative, and adaptive handover strategies.

H0₂ (Safety)
- There is no significant difference in perceived safety between the different handover strategies.

H0₃ (Trust)
- There is no significant difference in trust ratings across the three handover strategies.

# Task-2

### Understanding Power Analysis for t-tests (25%)

With the help of statsmodels1 library, create plots showing the relationship between alpha level, effect size, sample size and statistical power for both the cases described below. Along with this, please explain your observations based on the plots.

**Note:**
Take a look at the Python classes- *statsmodels.stats.power.TTestPower* and *statsmodels.stats.power.TTestIndPower*

*https://www.statsmodels.org/dev/stats.html#power-and-sample-size-calculations*

#### Case-1: Two independent sample t-test

<!-- grade_id: Task-2_A, points: 8, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

##### Explanation:

<!-- grade_id: Task-2_B, points: 4.5, grade: true, solution: true -->
YOUR ANSWER HERE

#### Case-2: Paired sample t-test

<!-- grade_id: Task-2_C, points: 8, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

##### Explanation:

<!-- grade_id: Task-2_D, points: 4.5, grade: true, solution: true -->
YOUR ANSWER HERE

# Task-3

### Performing statistical tests (30%)

A robotics company has created a new social companion robot for elderly persons. Since the robot is designed for elderly men and women who are living alone, it has to appeal to both demographic groups equally. Before the company mass-produces the robot, it wants to make sure that the robot appeals equally to both elderly men and women.


More specifically, the company wants to know whether men and women engage with the robot in the same way. To achieve this, 20 elderly men and 20 elderly women were recruited and asked to interact with the robot in a home setting. The participants were then asked to fill in a questionnaire that measures their engagement with the robot. An overall engagement score was computed based on the responses to the questionnaire. The scores for each person is provided in the file *‘advertising_agency_engagement.csv’*.

#### Plot appropriate graphs and comment on outliers, if any.

<!-- grade_id: Task-3_A, points: 6, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Determine whether the data is normally distributed.

<!-- grade_id: Task-3_B, points: 6, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Mention the statistical test that you would use to validate the hypothesis given below. Justify your answer.                       

***Hypothesis H1: Elderly men and elderly women engage differently with the robot.***

<!-- grade_id: Task-3_C, points: 6, grade: true, solution: true -->
YOUR ANSWER HERE

#### Conduct the statistical test and report the results.

<!-- grade_id: Task-3_D, points: 6, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Explain your conclusion about the hypothesis.

<!-- grade_id: Task-3_E, points: 6, grade: true, solution: true -->
YOUR ANSWER HERE

# Feedback

### Please answer the following:

#### How much time did you spend on doing this sheet per person? Anonymize your answer!

<!-- grade_id: Feedback_A, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

#### Was this sheet too easy / easy / ok / hard / too hard?

<!-- grade_id: Feedback_B, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

#### What additional resources (blogs, papers, books, tutorials, etc.) did you use? Please provide links or references.

<!-- grade_id: Feedback_C, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

#### Did you face any issues while solving this sheet?

<!-- grade_id: Feedback_D, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE