<div class="alert alert-block alert-info">
    
### Human-Centered Interaction and Collaboration in Robotics: Assignment 2

## General Information:
Please do not add or delete any cells. Answers belong into the corresponding cells (below the question). If a function is given (either as a signature or a full function), you should not change the name, arguments or return value of the function.

If you encounter empty cells underneath the answer that can not be edited, please ignore them, they are for testing purposes.

When editing an assignment there can be the case that there are variables in the kernel. To make sure your assignment works, please restart the kernel and run all cells before submitting (e.g. via Kernel -> Restart & Run All).

Code cells where you are supposed to give your answer often include the line `raise NotImplementedError`. This makes it easier to automatically grade answers. If you edit the cell please outcomment or delete this line.

## Submission:
Please submit your notebook via the web interface (in the main view -> Assignments -> Submit). For the deadline please refer to LEA.

## Group Work:
Please enter your UID (your username here) and those of your group partners into the next cell. We apply plagiarism checking, so do not submit others solutions! If an assignment has a copied solution, the task will be graded with 0 points for all people with the same solution.

## Questions about the Assignment:
If you have questions about the assignment please post them in the LEA forum ahead of the deadline. Don't wait until the last day to post questions!
    
</div>

### Please add the usernames of all your team members in the manner member1, member2 in next cell (example given below)

member1 = 'exampleone2s'

member2 = 'exampletwo2s'

member3 = 'examplefour2s'

#### If you are not working in a group or in a smaller group, then please add member3 and/or member2 as none2s

<!-- grade_id: cell-a9c2f4587c671f98, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e27ced94e387fe18, points: 0, grade: true, solution: false -->
```python
# Execute this cell to make sure you correctly filled in the usernames of the team members

def group_name_test():
    for member_id in [member1, member2, member3]:
        assert isinstance(member_id, str), "Please give your member id as a string."
        assert len(member_id) > 0, "You need to fill in the member id for both members"
        assert member_id.endswith("2s"), "The member id should end with 2s (Your JupyterHub username)"

group_name_test() 
print("All tests passed!")
```

# Ungraded-Task


This task is only for self study and will not be graded.

It is however recommended to answer these questions as it will act as practice for your exams.

#### Explain briefly the ethno-socio linguistic model for communication in humans.

<!-- grade_id: Ungraded-Task_A, points: 0, grade: true, solution: true -->
The ethno-Socio linguistic model for communication in human is based beyond sender-reciver models, it clarifies that communication context, where the meaning of a message can change depending on several parameters ;ike Temporal(When), Spatial(Where), Socio-culture(With whom) and external factors such as presence of ojects and audiences also influences. each individual has its personal set of goals to express and interpret communication signs. The core of this model is whether intentional or unintentional, we humans are always contineously communicating like even being silence in a particular context reflects communication. The model also talks about meta communication, which means communication about communication is also seen between conversation where temporarily orignal conversation is stop to seek more clarification on the communcation happened 

#### Name the components that constitute spoken dialog systems. What is the role of each of these components?

<!-- grade_id: Ungraded-Task_B, points: 0, grade: true, solution: true -->
ASR - Automatic Speech Recognition - this analyses acoustic signal from user and recognize the spoken words.

SLU - Spoken Language Understanding - This analyses the input form ASR to extract intents(querry, greeting, instruct etc) and entities(name , place , thing, time) 

DM - Dialog Management - The important part composed of Dialog Context Model(DCM), responsible for tracking the dialog state and information from NLU/SLU and Dialog Policy Model(DPL), responsible for making decision for next dialog act(intent,entities) considering the information given in DCM.

NLG - Natural Language Generation - it converst the dialog act into a sequence of Words

TTS - Text to Speech - it simply converts the words into the acoustic signals which are actually the response.

#### Explain the symbol grounding problem in robotics.

<!-- grade_id: Ungraded-Task_C, points: 0, grade: true, solution: true -->
symbol grouding is all about how the system abstracts information about a symbol to their actual meaning, in Robotics its like how in real workd setting the robot interpretates the word given to it considering ow the robot actually abstracts its information referring that immediate environemnt. its not about robot processing the data but about how robot understands the words meaningfully tied to its internbal representation of internal world 

# Task-1

### Verbal and Nonverbal Behaviours (30%)

#### Comment on the nonverbal communication capabilities of TIAGo Pro, NAO robot and aibo robot. Focus on kinesics (facial expressions, hand gestures, body movements, eye gaze) and vocalics, from the production and consumption points of view. (20%)

- Technical details of TIAGo Pro: https://pal-robotics.com/robots/tiago-pro/
- Technical details of NAO: http://doc.aldebaran.com/2-8/family/nao_technical/index_naov6.html
- Technical details of aibo: https://helpguide.sony.net/aibo/ers1000/v1/en-us/contents/TP0001970140.html

<!-- grade_id: Task-1_A, points: 20, grade: true, solution: true -->
YOUR ANSWER HERE

#### Represent the multimodal communicative behaviour shown in Figure 1 in the Behavior Markup Language format. (10%)

***Note:** You can assume that the human is standing in front of the robot and the robot only
needs to lift up its head by a specific amount to gaze at the face of the human. You are free to
set/tune the durations of each sub-behavior as you find appropriate/realistic.*

![Alt text](WS2_files/img/HCICR_A2.png)

<p style="text-align: center;">Figure 1: An example multimodal behavior.</p>

<!-- grade_id: Task-1_B, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

# Task-2

### Realise Multimodal Behaviour on Pepper (45%)

In this task, you will use the qiBullet simulation tool to realise the multimodal behaviour shown
in Figure 1.

Solution Note:

- You can use the threading feature in Python to start behaviours simultaneously, if necessary.
- Remember to reset the agent to the default state after the execution of the given behaviour episode.
- Please refer to the software requirements specification uploaded in LEA along with the homework.

Submission Note:

- Along with your solution here, please upload in LEA the source code and the screen recording showing the multimodal behaviour implemented on the Pepper robot in a zip archive.
- You do not need to insert and explain code snippets here, where needed, please specify the file location (in your LEA submission) which can be referred.

#### Specify the shown multimodal behaviour in the form of a dictionary in BML syntax. (10%)

<!-- grade_id: Task-2_A, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Extend your source code from Homework 1 to realise the behaviour shown in Figure 1. (30%)

<!-- grade_id: Task-2_B, points: 30, grade: true, solution: true -->
YOUR ANSWER HERE

#### Add comments in your code to explain your logic. (5%)

<!-- grade_id: Task-2_C, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Food for thought (ungraded):

Did you observe any challenges while synchronizing the different modalities? If yes, what information could have helped you to plan and schedule the behaviours more efficiently?

<!-- grade_id: Task-2_D, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

# Task-3

### Incremental Dialog Processing (25%)

Read the paper *(Skantze and Schlangen, 2009)* 

**Gabriel Skantze and David Schlangen. 2009. Incremental Dialogue Processing in a Micro- Domain. In Proceedings of the 12th Conference of the European Chapter of the ACL (EACL 2009), pages 745{753, Athens, Greece. Association for Computational Linguistics.**

and answer the questions given below:

#### What do you understand by "incremental dialog processing"? (5%)

<!-- grade_id: Task-3_A, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

#### How does incremental dialog processing support human-centeredness in interaction? (10%)

[Hint: See Section 2 in paper]

<!-- grade_id: Task-3_B, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

#### In the example presented in the paper, how did prosodic analysis contribute to incrementality? (5%)

<!-- grade_id: Task-3_C, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

#### In the architecture for spoken dialog systems presented in the paper, why is TTS connected to Discourse Modeller? (5%)

<!-- grade_id: Task-3_D, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

# Feedback

### Please answer the following:

#### How much time did you spend on doing this sheet per person? Anonymize your answer!

<!-- grade_id: Feedback_A, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

#### Was this sheet too easy / easy / ok / hard / too hard?

<!-- grade_id: Feedback_B, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

#### What additional resources (blogs, papers, books, tutorials, etc.) did you use? Please provide links or references.

<!-- grade_id: Feedback_C, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

#### Did you face any issues while solving this sheet?

<!-- grade_id: Feedback_D, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE