<div class="alert alert-block alert-info">
    
### Human-Centered Interaction and Collaboration in Robotics: Assignment 4

## General Information:
Please do not add or delete any cells. Answers belong into the corresponding cells (below the question). If a function is given (either as a signature or a full function), you should not change the name, arguments or return value of the function.

If you encounter empty cells underneath the answer that can not be edited, please ignore them, they are for testing purposes.

When editing an assignment there can be the case that there are variables in the kernel. To make sure your assignment works, please restart the kernel and run all cells before submitting (e.g. via Kernel -> Restart & Run All).

Code cells where you are supposed to give your answer often include the line `raise NotImplementedError`. This makes it easier to automatically grade answers. If you edit the cell please outcomment or delete this line.

## Submission:
Please submit your notebook via the web interface (in the main view -> Assignments -> Submit). For the deadline please refer to LEA.

## Group Work:
Please enter your UID (your username here) and those of your group partners into the next cell. We apply plagiarism checking, so do not submit others solutions! If an assignment has a copied solution, the task will be graded with 0 points for all people with the same solution.

## Questions about the Assignment:
If you have questions about the assignment please post them in the LEA forum ahead of the deadline. Don't wait until the last day to post questions!
    
</div>

### Please add the usernames of all your team members in the manner member1, member2 in next cell (example given below)

member1 = 'exampleone2s'

member2 = 'exampletwo2s'

member3 = 'examplefour2s'

#### If you are not working in a group or in a smaller group, then please add member3 and/or member2 as none2s

<!-- grade_id: cell-a9c2f4587c671f98, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e27ced94e387fe18, points: 0, grade: true, solution: false -->
```python
# Execute this cell to make sure you correctly filled in the usernames of the team members

def group_name_test():
    for member_id in [member1, member2, member3]:
        assert isinstance(member_id, str), "Please give your member id as a string."
        assert len(member_id) > 0, "You need to fill in the member id for both members"
        assert member_id.endswith("2s"), "The member id should end with 2s (Your JupyterHub username)"

group_name_test() 
print("All tests passed!")
```

# Task-1

### Human-Centered Reinforcement Learning (20%)

#### Name at least two characteristics of each of the three algorithms: TAMER, SABL and COACH. (10%)

<!-- grade_id: Task-1_A, points: 15, grade: true, solution: true -->
TAMER:
- Greedy Learning - it assumes when a user gives feedback , it has already factored it in the long term consequneces of action , so it only seeks maximum imediate effect rather than a discounted sum of future rewards.
- Supervised Learning - a learning function that allows how much to learn from feedback, where the state action pairs serves as the inputs and human scalar feedback serves as the labels.

SABL:
- Trainer Strategy Modeling - uses diferent strategies and combination (eg Positive reward , positive punishment , negative reward , negative punishment) that determins the feedback.
- Informative absense of Feedback - it understads the trainer strategy , can create lack of feedback as impict signal, like in case of reward focused trainer, a silent pause is statistically interpreted by the agent as impicit negative feedback.

COACH:
- Policy Dependent advanatge function - COACH is built on the psychological insights that human feedback changes based on what agent already knows, so it takes human feedback as an advantage function , measuring how better and worse could be the action related to agents current policy.
- Direct Policy Gradient updates - As human is a critic , COACH uses the feedback to make direct updates to the agents parameterized policy using eligibility traces. 


#### For each of the three algorithms, give an example of a human-robot interaction scenario where you would apply this algorithm. (10%)

<!-- grade_id: Task-1_B, points: 5, grade: true, solution: true -->
For TAMER , Considereing an example of teaching a robotic arm to sort fragile items , like the worker can gives direct evaluation to robot on its action creating a state action pair.

SABL can be used with an interactive social robot in a classroom, as in that environment childresna re unpredictable so the startegy can allow the robot to infer what kind of child it is interacting with , The hyper one or the one with patience.

COACH can e used in case of a case where a phycial apateint is recovering from a stroke , who has motorized lower-limb habilation exoskelton, here the patient and exosle;ton improve over weeks , the therapist naturally expects higher qulaity and stops prasing the basic movemnts , by giving feedback for highly refined, here this feedabck is treated as an advantage function , it understands the feedback is relative to the agents current level of mastery 

# Task-2

### TAMER for Mountain Car Problem (50%)

Here is a sample code for using TAMER to train a car to climb a mountain:

https://github.com/RitwikC07/Hill_climb_TAMER

Please refer to the following paper for the task:

Knox, W. & Stone, Peter. (2009). Interactively shaping agents via human reinforcement: the TAMER framework. 9-16.
10.1145/1597735.1597738.

*Submission Note:*

Along with your solution here, please upload in LEA the source code in a zip archive.
You do not need to insert and explain code snippets here, where needed, please specify the file location (in your LEA submission) which can be referred.

#### Modify the code base so that the program can take human evaluative feedback via keyboard with keys ‘W‘ or ‘w‘ for reward (+1), ‘A’ or ‘a‘ for punishment (-1), while the agent istraining. (5%)

<!-- grade_id: Task-2_A, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

#### Train the mountain car to reach the goal state when credit assignment is set to ‘False’ and share a screenshot of how far your car reached. (15%)

<!-- grade_id: Task-2_B, points: 15, grade: true, solution: true -->
YOUR ANSWER HERE

#### Explain the strategy you used for giving rewards and punishments while training the mountain car. (5%)

<!-- grade_id: Task-2_C, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

#### Repeat the last two steps with credit assignment now set to ‘True’. (20%)

<!-- grade_id: Task-2_D, points: 20, grade: true, solution: true -->
YOUR ANSWER HERE

#### Compare your results and reason about the role of credit assignment. (5%)

<!-- grade_id: Task-2_E, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

# Task-3

### Software architecture for your Chatbot project (30%)

#### Create a software architecture for the MAS Admission Assistance chatbot that you created. (30%) 

*You can refer to the slides for the conceptually complete architecture.*

<!-- grade_id: Task-3_A, points: 30, grade: true, solution: true -->
YOUR ANSWER HERE

# Task-4

### Ungraded Task (only for self-study)

#### With the help of a diagram illustrate and explain the processes and pathways in a "conceptually complete" multimodal architecture for socially interactive agents.

<!-- grade_id: Task-4_A, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

# Feedback

### Please answer the following:

#### How much time did you spend on doing this sheet per person? Anonymize your answer!

<!-- grade_id: Feedback_A, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

#### Was this sheet too easy / easy / ok / hard / too hard?

<!-- grade_id: Feedback_B, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

#### What additional resources (blogs, papers, books, tutorials, etc.) did you use? Please provide links or references.

<!-- grade_id: Feedback_C, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

#### Did you face any issues while solving this sheet?

<!-- grade_id: Feedback_D, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE